```python
# Initialize Otter
import otter
grader = otter.Notebook("TablePractice.ipynb")
```

Output:
```
File /opt/tljh/user/lib/python3.12/site-packages/otter/check/notebook.py:195, in Notebook._log_event(self, event_type, results, question, success, error, shelve_env)
    185         raise TypeError(f"shelve_env has an invalid type: {type(shelve_env)}")
    187     entry.shelve(
    188         shelve_env,
    189         delete=True,
   (...)
    192         variables=self._vars_to_store,
    193     )
--> 195 entry.flush_to_file(OTTER_LOG_FILENAME)

File /opt/tljh/user/lib/python3.12/site-packages/otter/check/logs.py:156, in LogEntry.flush_to_file(self, filename)
    153     dill.dump(self, file)
    155 except OSError:
--> 156     raise Exception(
    157         "Could not create the log file as the file system is read-only. Please contact your "
    158         "instructor before continuing on this assignment."
    159     )
    161 finally:
    162     if "file" in locals():

Exception: Could not create the log file as the file system is read-only. Please contact your instructor b
```

```python
import pandas as pd
import numpy as np
```

# 1. Citi Bike Dataset 🗽

Citibike is a well-known bike sharing/renting system in New York City. It is privately owned by Lyft and operates in the boroughs of the Bronx, Brooklyn, Manhattan, Queens, Jersey City New Jersey, and Hoboken New Jersey. The system was proposed in 2008 but officially started in 2013 with 332 stations and 6,000 bikes. As of today, there are 17,000 bikes with approximately 50,000 rides on average daily.

### Question 1.1

Write a python code to read the file **citibike-trips.csv** located under the data directory. Store your answer in the variable `citibike`.

Hint: Use `read_csv()` method

```python
citibike = ...
citibike.head()
```

```python
grader.check("q1.1")
```

### Question 1.2

Select the columns `tripduration`, `start_station_name` and `end_station_name` from the citibike dataframe. Store your answer in the updated dataframe `citibike`. 


```python
citibike = ...
citibike.head()
```

```python
grader.check("q1.2")
```

### Question 1.3

Find the average duration of a trip that lasted more than 30 minutes. Store your answer in the variable `above_30_mins` rounded upto two decimal points. 

```python
above_30_mins = ...
above_30_mins
```

```python
grader.check("q1.3")
```

### Question 1.4

Find the number of trips that started at the **W 78 St & Broadway** station. Store your answer in the variable `num_trips`. 


```python
num_trips = ...
num_trips
```

```python
grader.check("q1.4")
```

### Question 1.5

Find the station name where the most trips ended (Assume there are no ties). Store your answer in the variable `crowded_station` of type `str`.

```python
crowded_station = ...
crowded_station
```

```python
grader.check("q1.5")
```

### Question 1.6

Find the number of trips for which the trip duration ending at the `crowded_station` was at least 500 seconds. Store your answer in the variable `num_trips`. 

```python
num_trips = ...
num_trips
```

```python
grader.check("q1.6")
```

### Question 1.7

Find the number of stations for which the average duration ending at that station was at least 2000 seconds. Store your answer in the variable `num_stations`. 

```python
num_stations = ...
num_stations
```

```python
grader.check("q1.7")
```

# 2. SFA Basketball

The **sfa** table describes the `name` (string), `position` (string), `class` (string), and `height` (int) of SFA basketball
players in the 2024-125 season. Please note, this is just a synthetic data. 

```python
sfa = pd.read_csv('data/sfa_basketball.csv')
sfa.head()
```

### Question 2.1

Find the proportion of all players whose position is **Forward**. Store you answer in variable `fwd_prop` rounded upto two decimal points. 

```python
fwd_prop = ...
fwd_prop
```

```python
grader.check("q2.1")
```

### Question 2.2

Find the name of the shortest **freshman**. Assume that one is shorter than the rest. Store your answer in the variable `shortest` of type `str`


```python
shortest = ...
shortest
```

```python
grader.check("q2.2")
```

### Question 2.3

Find whether there are at least 3/4 of players that are atleast 80 inches taller. The result should be True or False and stored in the variable `percentile_75th`.

```python
percentile_75th = ...
percentile_75th
```

```python
grader.check("q2.3")
```

### Question 2.4

Find a pandas series of all positions, sorted in increasing order of the average height for all players in that position. Store your answer in variable `sorted_pos` of type `pandas.Series`


```python
sorted_pos = ...
sorted_pos 
```

```python
grader.check("q2.4")
```

Tests for q1.1 (points: None)
```
>>> citibike.shape == (473556, 16)
True
```


Tests for q1.2 (points: None)
```
>>> citibike.shape == (473556, 3)
True
```


Tests for q1.3 (points: None)
```
>>> int(above_30_mins) == 4117
True
```


Tests for q1.4 (points: None)
```
>>> num_trips == 362
True
```


Tests for q1.5 (points: None)
```
>>> crowded_station == 'E 17 St & Broadway'
True
```


Tests for q1.6 (points: None)
```
>>> num_trips == 2351
True
```


Tests for q1.7 (points: None)
```
>>> num_stations == 40
True
```


Tests for q2.1 (points: None)
```
>>> fwd_prop == 0.36
True
```


Tests for q2.2 (points: None)
```
>>> shortest == 'Lauryn Stockin'
True
```


Tests for q2.3 (points: None)
```
>>> percentile_75th == False
True
```


Tests for q2.4 (points: None)
```
>>> sorted_pos.iloc[0] == 'Guard' and sorted_pos.iloc[-1] == 'Center'
True
```
