```python
# Initialize Otter
import otter
grader = otter.Notebook("TablePractice2.ipynb")
```

```python
import pandas as pd
import numpy as np
```

## Give me some coffee

Jennifer is studying late at the SFA library and wants a coffee to help her stay alert and refreshed. She plans to order online so the drink will be ready for pickup when she drives to the coffee shop. She has access to two tables: `drinks` and `businesses`, that might help her select a coffee to order. 

`drinks` table:
- **Shop** (string): Name of the coffee shop
- **Price** (float): The base price of a coffee
- **Rating** (float): The aggregate Google reviews rating, out of 5.0 stars

`businesses` table: 
- **Business** (string): Name of the business
- **Neighborhood** (string): The neighborhood within Nacogdoches that the business is located in
- **Minutes** (integer): The estimated number of minutes that it takes to drive from SFA library to the business
- **Late_Night** (boolean): Indicates whether or not the business is open past 8 pm

```python
drinks = pd.read_csv('data/drinks.csv')
drinks.head(3)
```

```python
businesses = pd.read_csv('data/businesses.csv')
businesses.head(3)
```

### Question 1.1

Write a python code to create a pivot table displaying the number of coffee shops having a particular combination of neighborhood (as row labels) and late night availability (as column labels)


```python
your_ans = businesses.pivot_table(index = ... , columns=... , values=... , aggfunc=...)
your_ans
```

```python
grader.check("q1.1")
```

### Question 1.2

Find the number of coffee shops with a rating above 4.0 stars and a coffee prices less than 7 dollars.

```python
your_ans = ...
your_ans
```

```python
grader.check("q1.2")
```

### Question 1.3

Create a two column table containing the average price of coffee in shops for each neighborhood.


```python
merged_shops = ...
your_ans = ...
your_ans
```

```python
grader.check("q1.3")
```

### Question 1.4

The name of the restaurant which offers the best value for a coffee, given the amount of time it takes to drive to the restaurant from SFA library.

Hint: 
1. Find the Price per Minutes
2. Sort the columns based on Price_per_minutes
3. Find the coffee shop with least Price_per_minutes value


```python
merged_shops['Price_per_minutes'] = ...
sorted_merged_shops = ...
first_row = ...
your_ans = ...
your_ans
```

```python
grader.check("q1.4")
```

Tests for q1.1 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```
```
>>> your_ans.loc['North St', False].item() == 3
True
```
```
>>> your_ans.loc['North St', True].item() == 2
True
```


Tests for q1.2 (points: None)
```
>>> your_ans == 4
True
```


Tests for q1.3 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```


Tests for q1.4 (points: None)
```
>>> your_ans == 'Culinary Cafe'
True
```
